**Important note!** Before you turn in this lab notebook, make sure everything runs as expected:

- First, **restart the kernel** -- in the menubar, select Kernel$\rightarrow$Restart.
- Then **run all cells** -- in the menubar, select Cell$\rightarrow$Run All.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE."

# How is the average temperature in Atlanta changing over time?

One way to answer this question is to type it into Google. You might even get lucky and find that someone has answered it, though when I tried it (in preparing for this class), I did not find that to be the case: http://lmgtfy.com/?q=how+is+the+average+temperature+in+atlanta+changing+over+time

Here is another way: let's find the data and answer the question directly ourselves!

```python
import warnings
warnings.filterwarnings('ignore')

# For number crunching and data management:
import numpy as np
import pandas as pd
from IPython.display import display

# For plotting:
#import matplotlib.pyplot as plt
%matplotlib inline
import seaborn as sns

# For manipulating date/time stamps
import time
```

Start by downloading the data.

```python
!rm -f temperatures-us-cities.csv  # get a fresh copy
!curl -O -L https://github.com/rvuduc/cs1100/raw/master/temperatures-us-cities.csv
```

**Exercise 1.** Load the data into a Pandas data frame called `data` and inspect the first few (e.g., 10) rows.

<!-- grade_id: cell-35e22ddae619f9c2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Exercise 2.** Extract all the entries for the city of Atlanta. Store the result in a new data frame named `atl[:]` that contains just this subset of the original `data[:]`.

<!-- grade_id: atl, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Note the missing data, represented by `NaN` ("not-a-number") values. Lesson 1: In real life, you always need to do some kind of data postprocessing and cleaning. As an additional step, let's drop any rows that have missing data.

```python
?atl.dropna
```

```python
atl.dropna (inplace=True)
atl.head (10)
```

**Exercise 3.** Create a new column named 'year' containing just the year (as an integer).

<!-- grade_id: get_year, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Exercise 4.** Group the data by year and compute the mean temperature and uncertainty. Store the result in a data frame named `atl_means`.

<!-- grade_id: atl_means, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Exercise 5.** Convert the degrees Centigrade values to Farenheit. Store these values in two columns: `degF` (for `AverageTemperatureC`) and `deltaF` (for `AverageTemperatureUncertainty`).

<!-- grade_id: degC_to_degF, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Now let's plot the data to answer the question.

```python
sns.regplot (x='year', y='degF', data=atl_means)
```

**Discussion.** There appears to be a warming trend though there is also a fair amount of variability. As such, what would you say are the limits to this analysis? Some suggestions for further exploration:

- What do you make of the "uncertainty" field (computed as `atl_means['deltaF']`)?
- Does the trend hold in other cities?
- What sort of follow-up questions, analyses, data collection exercises, or simulations do you think would be useful for explaining _why_ this trend exists?

<!-- grade_id: final_test, points: 1, grade: true, solution: false -->
```python
import statsmodels.api as sm
y = atl_means['degF'] # Response
x = atl_means['year'] # Predictor
x = sm.add_constant(x)
model = sm.OLS(y, x)
results = model.fit()
delta_year = x['year'].max() - x['year'].min()
delta_F = results.params['year']*delta_year
print("==> Over {} years, the temperature changed by about {:.2f} degrees Farenheit.".format(delta_year, delta_F))
assert 2.48 <= delta_F <= 2.49
print("\n(Passed!)")
```

```python

```